---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
---

## Introducción

```yaml
type: NormalExercise
key: 2bafef99a3
lang: r
xp: 100
skills: 1
```

La estadística permite que......

`@instructions`
# esta prueba

`@hint`
[prueba]()

`@pre_exercise_code`
```{r}
a=10
```

`@sample_code`
```{r}
a=15
```

`@solution`
```{r}
a=20
```

`@sct`
```{r}

```

---

## kkkk

```yaml
type: MultipleChoiceExercise
key: 0a92db9ab5
xp: 50
```



`@possible_answers`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sct`
```{r}

```
