---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
free_preview: true
---

## Introducción

```yaml
type: NormalExercise
key: 2bafef99a3
lang: r
xp: 100
skills: 1
```

La estadística permite que......

`@instructions`
# esta prueba

`@hint`
[prueba]()

`@pre_exercise_code`
```{r}
a=10
```

`@sample_code`
```{r}
a=15
```

`@solution`
```{r}
a=20
```

`@sct`
```{r}

```
